# Homework 3


## Problem 1 (Reocurring Characters: dictionaries) 

For a given string s, return the first reoccuring character in the string.  Try to get an approach that is $O(n)$, where $n$ is the number of characters in the string. If there is no such character, return None. 
    
    -- Input:"ABCA"
    -- Output:"A"

```python
def problem1(s):
    occurent_letter = None
    final_dict ={}
    # YOUR CODE HERE
    for str in s:
        if str not in final_dict:
            final_dict[str] = 1
        else:
            occurent_letter = str
            return(occurent_letter)
# s ="ABCA"
# problem1(s)
```

<!-- grade_id: cell-2853176972493b6e, points: 10, grade: true, solution: false -->
```python
assert problem1("ABCA") == "A"
```

<!-- grade_id: cell-b16dd3056b5a9f95, points: 10, grade: true, solution: false -->
```python
assert problem1("ABC") == None
```

<!-- grade_id: cell-ab1b1d0a5f7b1fce, points: 10, grade: true, solution: false -->
```python
assert problem1("ABCBA") == "B"
```

<!-- grade_id: cell-b78086c290d4a60c, points: 0, grade: false, solution: true -->
```python
#Write ChatGPT code here
def first_reoccuring(s):
    seen = set()
    for char in s:
        if char in seen:
            return char
        seen.add(char)
    return None

# Test
s = "ABCA"
print(first_reoccuring(s))  # Outputs: A

```

Output:
```
A
```

## Problem 2 (String Methods):

We have learnt that strings are immutable objects. In the following, you are asked to write a code to "mutate" a string. In particular, you are given a string, a position number and a character. You should return a different string such that this string is the same as the input string except the character in the position is changed to the given character. 

Example:

    input: 
        original_string = "Hello world!"
        position=6
        character="W"
    ouput: "Hello World!"

```python
def problem2(original_string, position, character):
    return_string = ""
    
    # YOUR CODE HERE
    return_string = original_string[:position] + character + original_string[position+1:]
    
    return(return_string)
```

<!-- grade_id: cell-681343a4d7e22bbd, points: 10, grade: true, solution: false -->
```python
# Test cases for Problem Set 4
assert problem2("Hello world!", 6, "W") == "Hello World!"
```

<!-- grade_id: cell-402bf89f76c4bd42, points: 10, grade: true, solution: false -->
```python
assert problem2("Hello world!", 11, ".") == "Hello world."
```

<!-- grade_id: cell-562beecb93f03f23, points: 0, grade: false, solution: true -->
```python
#Write ChatGPT code here (same as mine)
def mutate_string(original_string, position, character):
    # Use slicing to create a new string with the desired modification
    return original_string[:position] + character + original_string[position+1:]

original_string = "Hello world!"
position = 6
character = "W"
print(mutate_string(original_string, position, character))  # Outputs: Hello World!

```

Output:
```
Hello World!
```

## Problem 3 (Function)

Write a function that checks whether two words are anagrams. Two words are anagrams if they contain the same letters. For example, silent and listen are anagrams, dog and god are anagrams.

Your code must be able to deal with follwoing three cases:

Input: "silent", "listen"

output: silent and listen are anagrams

        1
        
Input: "dog", "good"

output: The two strings do not have the same lenght

        -1

Input: "cat", "tag"

output: cat and tag are not anagrams

        0

```python
def problem3(s1, s2):
    if sorted(s1)==sorted(s2):
        print(s1 + " and " + s2 + " are anagrams")
        return 1
    elif len(s1) != len(s2):
        print("The two strings do not have the same length")
        return -1
    else:
        print(s1 + " and " + s2 + " are not anagrams")
        return 0
    
    #Write your code here
```

```python
assert problem3("silent", "listen") == 1
```

Output:
```
silent and listen are anagrams
```

```python
assert problem3("dog", "good") == -1
```

Output:
```
The two strings do not have the same length
```

```python
assert problem3("create", "regret") == 0
```

Output:
```
create and regret are not anagrams
```

```python
problem3("silent", "listen")
```

Output:
```
silent and listen are anagrams
1
```

```python
problem3("dog", "good")
```

Output:
```
The two strings do not have the same length
-1
```

```python
problem3("create", "regret")
```

Output:
```
create and regret are not anagrams
0
```

```python
problem3("cat", "act")
```

Output:
```
cat and act are anagrams
1
```

```python
#Write ChatGPT code here
```

```python
def are_anagrams(word1, word2):
    # Check for same length
    if len(word1) != len(word2):
        return f"The two strings do not have the same length\n-1"

    # Check if sorted versions of the words are the same
    if sorted(word1) == sorted(word2):
        return f"{word1} and {word2} are anagrams\n1"
    else:
        return f"{word1} and {word2} are not anagrams\n0"

# Test cases
print(are_anagrams("silent", "listen"))
print(are_anagrams("dog", "good"))
print(are_anagrams("cat", "tag"))

```

Output:
```
silent and listen are anagrams
1
The two strings do not have the same length
-1
cat and tag are not anagrams
0
```

```python

```